<?xml version="1.0" encoding="UTF-8"?>


<listing>
    <caption>A Python program with preamble/postamble</caption>
    <program interactive='activecode' language="python" label="program-activecode-python-ambles" include-source="yes">
        <preamble>
        def add(a, b):
        </preamble>
        <code>
            # TODO - complete the add function

        </code>
        <postamble>
        # Use the function
        result = add(2, 3)
        if result == 5:
            print("Test passed")
        else:
            print("Test failed")
        </postamble>
    </program>
</listing>


<program interactive="activecode" language="python">
      <code>
      def add(a,b):
          return 4
      </code>
      <tests>
      from unittest.gui import TestCaseGui

      class myTests(TestCaseGui):

         def testOne(self):
             self.assertEqual(add(2,2), 4, "Adding two identical integers")
             self.assertEqual(add(7,13), 20, "Adding two different integers")
             self.assertAlmostEqual(add(2.0,3.0), 5.0, 5, "Adding two decimal numbers")

      myTests().main()
      </tests>
  </program>


  <exercise label="fillin-integer">
            <title>Fill-In, Integer Answer</title>

            <statement>
                <p>The game of bowling uses <var width="28"/> pins that you try to knock down. (This answer blank has been set to be very wide.)</p>
            </statement>
            <setup>
                <var>
                    <condition number="10">
                        <feedback>
                            <p>Arranged in a triangle, there are <m>1+2+3+4 = 10</m> pins, a so-called <term>triangular</term> number.</p>
                        </feedback>
                    </condition>
                    <condition number="16">
                        <feedback>
                            <p><em>Close</em>! You may have used hexadecimal notation, when you did not really mean to.</p>
                        </feedback>
                    </condition>
                    <condition string=".*">
                        <feedback>
                            <p>Incorrect.</p>
                        </feedback>
                    </condition>
                </var>
            </setup>
        </exercise>